# PSET 0: C++ Review & Matrix Multiplication

This problem set is designed to refresh your C++ skills and then get your matrix chops warmed up via a simple multiplication implementation. You will work through problems that cover C++ basics, arrays, dynamic memory, and templates. Each C++ section builds on the previous one, so make sure to work through them in order.

You can compile and run your C++ code within this notebook. Follow the instructions provided in each section.  We will run autograder tests on your work (you should be able to see the results in Gradescope).  For a few problems, the autograder tests are embedded right in your notebook and you can run them yourself!

Try not to add additional cells nor rearrange any that exist now as it will mess with our autograder (but feel free to open another Colab/notebook on the side).

### Make sure to submit your final notebook with all of your solutions to Gradescope!
[Direct Gradescope Link](https://www.gradescope.com/courses/820552)

### Before Starting

We're in an interactive _Python_ notebook.  So, of course, running C++ code is not going to be as easy as it is to run Python code.  That said, we can use a helper (just a Python wrapper to ingest and spit out a `.cpp` file) to make this cleaner.  Start by importing it below.

```python
# Install some magic to make c programs look nice!
!curl -o ./cpp_saver.py https://raw.githubusercontent.com/COMS-BC3159-F24/helpers/main/cpp_saver.py
%load_ext cpp_saver
```

#### Navigating
Nice, now that we've installed this plugin, let's expore where things will go.  We can use the command line to view all files in the current working directory by using the `ls` command. We can also use the `ls -al` command to view all files (including hidden ones) along with additonal information about each file. Note: in Google Colab all command line commands need to be preceeded by a `!` or through a special function like `%%writefile` as shown above.

```python
!ls -al
```

#### Using the Plugin
To use the plugin, we'll simply add a `%%cpp -n [your_filename].c` line to the start of each code block in our interactive notebook.  Then we can write our `C` or `C++` as we wish!  This will take the C++ you write below in the cell and save it to that file.  Let's try it for a simple `Hello, World!`

```python
%%cpp -n hello.cpp

#include <cstdio>
int main() {
    // C and C++ programs start executing from the main function
    // the main function returns an integer (it needs to be "typed")
    printf("Hello, Barnard!\n");
    return 0;
}
```

You'll see that we have successfully made the `hello.cpp` file in the current folder (which is referred to as `.`). We can also create folders by using the command `mkdir` and move files into that folder with the `mv` command (the `cp` command works similarly and copies files).  You can change directories (move your current position from folder to folder) using `cd` but this sort of stuff starts to break down in our Colab/`.ipynb` world, so we won't do much of this.  All of the `.cpp` files and executables we generate will get dumped in the same directory as your `.ipynb` file.

#### Compiling and Running
We've captured our input to a file, but let's actually compile and run!  Let's now compile and run our code.  As a compiled language you need to first run a compiler over an `C` code to produce binary code that the computer can then run. While this two step process may seem cumbersome to those of you coming from a mostly Python background, this allows for additional optimizations for runtime. Below is a figure from Dive into Systems that shows this visually!

![compile](https://diveintosystems.org/singlepage/book/modules/C_intro/assets/images/compile.png)

```python
# Now lets compile our code into a working file!
!g++ hello.cpp -o hello
```

At this stage, you might find tons of pesky errors in your code and cryptic compilation errors.  Start at the top of the output and work on each line of code that the compiler error mentions until you fix all errors!  Google and StackOverflow are always useful for this!

```python
# And now lets run that code!
!./hello
```

## Problem 1: C++ Basics
**Note**: our autograder requires that your output matches exactly the form we provide as example

### Part A: Basic Syntax and I/O
Write a simple C++ program that copies and inputted integer, doubles the copy, and prints the original and doubled version using `printf()` (since CUDA and `std::cout` don't play nicely) like so:

```text
Original: 1
Doubled: 2
```
This will become very handy for quick CUDA debugging. For more information on `printf()` and format codes like `%.1f` and `%d` and special characters like `\n` please see: https://en.cppreference.com/w/c/io/fprintf.

*Hint: those codes/characters just might be useful in your answer!*

### Part B: Control Structures
Extend your program to check if the number is `positive`, `negative`, or `zero`, and print a message exactly as follows like so but with the respective pieces for a given input:

```text
Original: 1
Doubled: 2
Original is: positive
```


```python
%%cpp -n basic_io.cpp

#include <cstdio>       // for printf
#include <string>       // for std::stoi

int main(int argc, char *argv[]) {

    // Check if the correct number of arguments are provided
    if (argc < 2) {
        printf("Usage: %s <number>\n", argv[0]);
        return 1;
    }

    // Convert the first argument to an integer
    int input_val = std::stoi(argv[1]);

    /* Part A: Basic Syntax and I/O */
    // Double the number
    // TODO
    ...
    ...
    // Print original and new
    // TODO
    ...
    ...

    /* Part B: Control Structures */
    // Print the sign
    // TODO
    ...
    ...
    ...
    ...
    ...
    ...
    ...

    return 0;
}
```

#### Let's compile and run it
Hopefully, you get the same output as we suggested.  You might have to deal with those pesky errors, first!

```python
!g++ basic_io.cpp -o basic_io
!./basic_io 1
```


## Problem 2: Arrays, Looping, and Indexing

Write a program that initializes an array of 10 integers, fills it with the numbers `n` through `n+10`, where `n` was provided as input. Then calculate the sum of the array elements.

Print the array by calling our `print_array()` function and then print the sum, so your output should look **exactly** like the following for an input of starting integer 1:

```text
Array elements: 1, 2, 3, 4, 5, 6, 7, 8, 9, 10,
Sum: 55
```


```python
%%cpp -n arrays.cpp

#include <cstdio>       // for printf
#include <string>       // for std::stoi

// Another function for printing
void print_array(int arr[], int size) {
    printf("Array elements: ");
    for (int i = 0; i < size; i++) {
        printf("%d, ", arr[i]);
    }
    printf("\n");
}

int main(int argc, char *argv[]) {

    // Check if the correct number of arguments are provided
    if (argc < 2) {
        printf("Usage: %s <number>\n", argv[0]);
        return 1;
    }

    // Convert the first argument to an integer
    int input_val = std::stoi(argv[1]);

    int arr[10];
    int sum = 0;

    // Fill the array with the specified values and print the array
    // TODO
    ...
    ...
    ...
    ...
    ...
    // Calculate the sum of array elements and print the sum
    // TODO
    ...
    ...
    ...
    ...

    return 0;
}
```

#### Compile and Run
Do you get the same output as we suggested?  If not, revisit your implementation!

```python
!g++ arrays.cpp -o arrays
!./arrays 1
```


## Problem 3: Dynamic Memory with Pointers

Repeat Problem 2 but with a program that *dynamically* allocates memory for an array of integers and deallocates.

Ensure your program handles potential memory leaks or dangling pointers.


#### An Aside on Pointers

In C we get to directly modify memory addresses. This allows us significantly more control over data movement which we will leverage in this course for acceleration. However, it also can lead to lots of security vulnerabilities and so we'll need to be careful when using pointers.

Ok so what is a pointer? **A pointer is a special kind of variable which holds a memory address!** Aka it "points" at the place in memeory where some value is stored. It turns out that is just what object variables are in Java and Python but in C we are able to treat them simply as numbers and manipulate them a lot more! This also means we can simply pass compact numbers from one function to another instead of all of the data! Below is an image from Dive into Systems that shows this graphically.

![pointer](https://diveintosystems.org/singlepage/book/modules/C_depth/assets/images/ptr.png)

Programs often allocate memory dynamically to tailor the size of an array for a particular run.  They "point" to that allocation via a "pointer."

We allocate memory using
```cpp
int* my_pointer = (int*) malloc(sizeof(int))
```
You may need to increase the size passed to `malloc()` if you're allocating more than a single integer.  And you may need to change the type of pointer and the casting (just preceding `malloc()`) if you're allocating space for a different type.

For good practice, you might want to check if the pointer is `nullptr` and error if so.

When you're done using the dynamically allocated memory, be sure to call `free(my_pointer)`.  (It's important to keep track of the memory that has been allocated dynamically, as forgetting to deallocate it can cause memory leaks and may lead to other memory-related problems in the future.)

#### Problem 3 (repeated)

Write a program that allocates space for an array of 10 integers, then fill it with the numbers `n` through `n+10`, where `n` was provided as input. Then calculate the sum of the array elements.

Print the array by calling our `print_array()` function and then print the sum, so your output should look **exactly** like the following for an input of starting integer 1:

```text
Array elements: 1, 2, 3, 4, 5, 6, 7, 8, 9, 10,
Sum: 55
```

```python
%%cpp -n dynamic.cpp

#include <cstdio>       // for printf
#include <string>       // for std::stoi

void print_array(int* arr, int size) {
    printf("Array elements: ");
    for (int i = 0; i < size; i++) {
        printf("%d, ", arr[i]);
    }
    printf("\n");
}

int main(int argc, char *argv[]) {

    // Check if the correct number of arguments are provided
    if (argc < 2) {
        printf("Usage: %s <number>\n", argv[0]);
        return 1;
    }

    // Convert the first argument to an integer
    int input_val = std::stoi(argv[1]);
    int arr_size = 10;

    // Dynamically allocate memory for the array
    // TODO
    ...
    ...
    // Fill the array with the specified values and print the array
    // TODO
    ...
    ...
    ...
    ...
    ...
    ...
    // Calculate the sum of array elements and print the sum
    // TODO
    ...
    ...
    ...
    ...
    ...
    ...
    ...
    ...
    ...

    return 0;
}
```

#### Compile and Run
Do you get the same output as we suggested?  If not, revisit your implementation!

```python
!g++ dynamic.cpp -o dynamic
!./dynamic 1
```


## Problem 4: Templates

Some of the difficulty of the "strict-typed" nature of `C++` can be eased by using "templates."  Templates allow you to define generic classes (blueprints for more complicated objects), functions, and more without specifying a type.  Then when you use the templated function, for example, you can explicitly set the type or let the compiler infer it.  This allows for writing generic code without having to copy-and-paste lots of implementations.

Take, for example, a simple classroom example where I want to write a function that adds one to a value.  Unlike Python which would happily let me define such a function without specifying the input type, in `C++`, I would need to be explicit:

```cpp
int addOne(int a) {
    return a + 1;
}

float addOne(float a) {
    return a + 1.0f;
}

long addOne(long a) {
    return a + 1;
}

double addOne(double a) {
    return a + 1.0;
}
```

This is obviously tedious, prone to error, and makes it hard to change code because I'd need to change 4 places.  Instead, I could use a template for both the argument and return type, and then when I call the function, the types get set in place:

```cpp
template <typename T>
T addOne(T a) {
    return a + 1
}

int main(void) {
    long long my_numb = 3159;
    long long result = addOne(my_numb);
    // or instead of relying on automatic template deduction
    // you could call addOne<long long>(my_numb);

    printf("My value + 1: %lld\n", result);
}
```

#### Your Turn!

Create a template function for `max()` that takes two arguments of any type and returns the larger of the two. Feel free to test your function with different data types (e.g., integers, doubles, and strings) and just use local variables that you set yourself and comment our code that handles the commandline args.  

When you're done, put back our commandline argument handling code and we'll test exactly these inputs and expect exactly the output shown (be careful with the number of decimal places for your floating point output):

```text
./templates float 4.5 3.1
Max of 4.5 and 3.1 is 4.5

./templates int 1 10
Max of 1 and 10 is 10
```

```python
%%cpp -n templates.cpp

#include <cstdio>
#include <string>

...
...
...
...

int main(int argc, char* argv[]) {
    if (argc != 4) {
        printf("Usage: %s <type> <value1> <value2>\n", argv[0]);
        return 1;
    }

    std::string type = argv[1];
    std::string val1 = argv[2];
    std::string val2 = argv[3];

    if (type == "int") {
        int a = std::stoi(val1);
        int b = std::stoi(val2);
        // TODO: Compute and print the integer case!
        ...
    }
    else if (type == "float") {
        float a = std::stof(val1);
        float b = std::stof(val2);
        // TODO: Compute and print the float case!
        ...
    }
    else {
        printf("Unsupported type: %s\n", type.c_str());
        return 1;
    }

    return 0;
}

```

#### Compile and Run
Do you get the same output as we suggested?  If not, revisit your implementation!

```python
!g++ templates.cpp -o templates
!./templates float 4.5 3.1
!./templates int 1 10
```

---


## Problem 5: Python Matrix Multiplication

In class, we've been discussing the performance of various matrix operations and the possibility of using parallelism to improve our efficiency (loosely defined).  The next problem set coding assignments will explore this in greater detail and with a slew of benchmarking, but for now, let's remind ourselves how to code the naïve algorithm for matrix multiplication but using Python (as it's easier to start this way than in C/C++).

In the written part of PS0, you likely reacquainted yourself with the standard matrix multiplication algorithm:
$$C_{ij} = \sum_{k=1}^{n} A_{ik} \times B_{kj}$$

We can apply this to a simple example for the purposes of starting our lab.  Consider a matrix $A$ (in gray) which is $4\times5$ and $B$ (blue) of dimensions $5\times4$ whose product results in $C$ (purple) of size $4\times4$.

![matrix mult example](https://raw.githubusercontent.com/COMS-BC3159-F24/image_assets/main/matmul_example.png)

We've provided the initial matrices (which we're implementing using `list`s of `list`s).  And, we've also provided a handy `printMatrix()` routine to do just what it says, print a matrix.

```python
def printMatrix(matrix):
    for row in matrix:
        print(' '.join(f'{elem:5}' for elem in row))

A = [
        [0, 1, 2, 3, 4],
        [5, 6, 7, 8, 9],
        [10, 11, 12, 13, 14],
        [15, 16, 17, 18, 19]
    ]

B = [
        [20, 21, 22, 23],
        [24, 25, 26, 27],
        [28, 29, 30, 31],
        [32, 33, 34, 35],
        [36, 37, 38, 39]
    ]
```

```python
def matrix_multiply(A, B):
    ...
    return C

C = matrix_multiply(A, B)
printMatrix(C)
```

Tests for q1-basics (points: 2)


Tests for q2-arrays (points: 2)


Tests for q3-dynamic (points: 2)


Tests for q4-templates (points: 2)
```
>>> import random
>>> import subprocess
>>> 
>>> def test_cpp_output(cpp_type, random_number_one, random_number_two):
...     result = subprocess.run(['./templates', f'{cpp_type}', f'{random_number_one}', f'{random_number_two}'], capture_output=True, text=True)
...     output = result.stdout.strip().split('\n')
...     if cpp_type == 'float':
...         solstring = f'Max of {random_number_one:.1f} and {random_number_two:.1f} is {max(random_number_one, random_number_two):.1f}'
...     else:
...         solstring = f'Max of {random_number_one} and {random_number_two} is {max(random_number_one, random_number_two)}'
...     assert output[0] == solstring
>>> test_cpp_output('int', random.randint(1, 10), random.randint(1, 10))
>>> test_cpp_output('float', float(random.uniform(1.0, 10.0)), float(random.uniform(1.0, 10.0)))
```


Tests for q5-matmul (points: 2)
